# ECON-GA-4002: Homework 2

**Name**: Jia Yi Kung



## Part 1: Theoretical Foundations

The questions in this section from `01_theory.ipynb`

Please review that notebook for context and help

<a id='exercise-0'></a>
**Exercise 1.1**

Show that for a convex set $C \subseteq \mathbb{R}^n$ any convex combination of members of $C$ is in $C$

<!-- grade_id: ex1_1, points: 10, grade: true, solution: true -->
Assume that two points x1, x2 belongs to the convex set C, then their combination will be Î»*x1+(1-Î»)*x2. Since we could find out that thier combination is just the line segment [x1, x2], we can conclude that any convex combination of members of C is in C.

<a id='exercise-2'></a>
**Exercise 1.2**

Show that if $C$ is convex, then $\textbf{conv} C = C$

<!-- grade_id: ex1_2, points: 10, grade: true, solution: true -->
From the definition of the convex hull of a set C (convC), we could know that convC is the set of all convex combinations of points in C. Further, As we know in Exercise 1.1, any convex combination of members of C is in C, we can conclude that convC contains C and convC is a convex set. 

Therefore, all convex sets containing C must must contain convC, and convC is itself a convex set containing C, convC = C.

### Convex Functions

A function $f: \mathbb{R}^n \rightarrow \mathbb{R}$ is *convex* if :

1. $\textbf{dom} f$ is convex
2. $\forall x, y \in \textbf{dom} f$ and $\theta \in [0, 1]$ $$f(\theta x + (1 - \theta) y) \le \theta f(x) + (1 - \theta) f(y)$$


<a id='exercise-3'></a>
**Exercise 1.3**

Show that all linear functions are convex

*Hint:* To do this show that the convexity inequality above (condition 2) holds with equality when $f$ is linear

<!-- grade_id: ex1_3, points: 10, grade: true, solution: true -->
To prove that all linear functions are convex, we need to prove that every line segment joining two points on the convex curve never lies below the curve at any point. 

Assume f be a linear function defined on the interval [x1, x2], and x1 â‰¤ a â‰¤ x2, x1 â‰¤ b â‰¤ x2, a < b. From the convex curve, we then could find out the height of the line segment (a, f(a)) to (b, f(b)) â‰¥ f(x). Hence, we can conclude: when f is linear, for any value of Î» with 0 â‰¤ Î» â‰¤ 1, ð‘“(ðœƒa+(1âˆ’ðœƒ)b)â‰¤ðœƒð‘“(a)+(1âˆ’ðœƒ)ð‘“(b).

# Part 2: Linear programs

The questions in this section from `02_linear_programs.ipynb`

Please review that notebook for context and help

<a id='exercise-0'></a>
**Exercise 2.1**

Convert a linear program in the general form to a linear program in standard form


*Hint:* You will need to introduce a new vector (of slack variables) into the constraints

<!-- grade_id: ex2_1, points: 10, grade: true, solution: true -->
Assume we need to minimize 80x+60y, subject to x+y â‰¥ 1, -0.05x+0.07y â‰¤ 0, x, y â‰¥ 0. 

To convert to standard form, we introduce two new variables s1 â‰¥ 0, s2 â‰¥ 0. The first measures how much over 1 the quantity x + y is, and the second measures how much under 0 the quantity âˆ’.05x + .07y is. 

Therefore, we could get: minimize 80x+60y, subject to x+y-s1 = 1, -0.05x+0.07y+s2 = 0, x, y, s1, s2 â‰¥ 0.

**Exercise 2.2: Transportation problem**

> Note: This problem comes from "Labs for Foundations of Applied Mathematics Volume 2" by Jeffrey Humpherys And Tyler J. Jarvis released under the [Creative Commons Attribution 3.0 United States license](http://creativecommons.org/licenses/by/3.0/us/). The original source for the lab materials can be found at [https://github.com/Foundations-of-Applied-Mathematics/Labs](https://github.com/Foundations-of-Applied-Mathematics/Labs)

Consider the following transportation problem: 

- A piano company needs to transport thirteen pianos from their three supply centers (denoted by 1, 2, 3) to two demand centers (4, 5)
- Transporting a piano from a supply center to a demand center incurs a cost (see table below)
- The company wants to minimize shipping costs for the pianos while meeting the demand

| Supply Center | Demand Center | Cost of Transportation | Number of Pianos | 
| :-----------: | :-----------: | :--------------------: | :--------------: |
| 1 | 4 | 4 | $p_1$ |
| 1 | 5 | 7 | $p_2$ | 
| 2 | 4 | 6 | $p_3$ | 
| 2 | 5 | 8 | $p_4$ |
| 3 | 4 | 8 | $p_5$ |
| 3 | 5 | 9 | $p_6$ |
The number of pianos available at each supply center is given by:

| Supply Center | Number of pianos available |
| :-----------: | :------------------------: |
| 1 | 7 |
| 2 | 2 | 
| 3 | 4 |
The number of pianos needed at each demand center is given by

| Demand Center | Number of pianos needed |
| :-----------: | :---------------------: |
| 4 | 5 |
| 5 | 8 |
A linear program can be set up to solve for $p_i$, $i=1,\dots,6$

The objective is to minimize total cost of shipping pianos, subject to these constraints:

- There cannot be a negative number of pianos transported along any route
- All supply centers must transport all pianos
- All demand centers must end up with needed number of pianosFinally, the objective function is the number of pianos shipped along each
route multiplied costs found in the costs table.
Your task is to formulate the transportation problem as a linear program

That is define the vector $c$, matrix $A$ and vector $b$ that appropriately defines the problem as described above

*Hint:* the matrix $A$ will have 5 rows and 6 columns

<!-- grade_id: 236c40, points: 0, grade: false, solution: true -->
```python
import numpy as np

def piano_transporation_problem():
    """
    Define the arrays c, A, and b that set up a linear program
    for solving the optimal piano transport problem described above

    Returns
    -------
    c: np.array
        The coefficient vector c

    A: np.array
        The constraint matrix A

    b: np.array
        The constraint value vector b

    Notes
    -----

    In order to pass the tests below, you must have the rows of A 
    and b match order of location numbers from the tables above

    """
    c = np.array([4., 7., 6., 8., 8., 9.])
    A = np.array([
        [1,1,0,0,0,0],
        [0,0,1,1,0,0],
        [0,0,0,0,1,1],
        [-1,0,-1,0,-1,0],
        [0,-1,0,-1,0,-1]
    ])

    b = np.array([7., 2., 4., -5., -8.])

    return c, A, b
```

<!-- grade_id: ex2_2, points: 20, grade: true, solution: false -->
```python
from scipy.optimize import linprog

c, A, b = piano_transporation_problem()

assert np.allclose(A.sum(axis=1), [2, 2, 2, -3, -3])
assert np.allclose(A.sum(axis=0), np.zeros(6))
assert all(c > 0)

# solve with linprog and check solution
sol = linprog(c, A, b)
assert np.allclose(sol.x, [5, 2, 0, 2, 0, 4])
```

# Part 3: Simplex Algorithm

The questions in this section from `03_simplex_algorithm.ipynb`

Please review that notebook for context and help

**Exercise 3.1**

Write a Python function that can consume `cbar`, `Abar` and `b` and return the tableau $D$:

<!-- grade_id: b36601, points: 0, grade: false, solution: true -->
```python
import numpy as np

def build_tableau(cbar: np.ndarray, Abar: np.ndarray, b: np.ndarray) -> np.ndarray:
    """
    Given cbar, Abar, and b for a stanard linear program
    
    construct a tableau of the form
    
    D = [[0, cbar.T], [b - Abar]]
    """
    m, n = Abar.shape
    assert (len(b), len(cbar)) == (m,n)
    up=np.hstack((np.array([0]), cbar))
    down=np.column_stack((b, -Abar))
    D=np.vstack((up, down))
    return D
```

<!-- grade_id: ex3_1, points: 10, grade: true, solution: false -->
```python
# from 01_simplex_algorithm.ipynb
def to_standard_form(c, A, b):
    """
    Convert a linear program of the form
    
    min_x dot(c, x) s.t. Ax <= b and x >= 0
    
    into a standrd linear program of the form
    
    min_xbar dot(cbar, xbar) s.t. Ax + s = b and xbar >= 0
    """
    # check dimensions
    m, n = A.shape
    assert (len(b), len(c)) == (m, n)
    cbar = np.concatenate([c, np.zeros(m)])
    Abar = np.column_stack([A, np.eye(m)])
    
    # b is unchanged
    return cbar, Abar, b

# product mix problem from 02_linear_programs.ipynb
p = np.array([250., 215., 275., 180.])  # price vector
d = np.array([10., 20., 12., 10.])      # demand vector
m = np.array([4., 4., 4.])              # resource constraints

H = np.array([                          # resource usage
    [0.12, 0.18, 0.13, 0.07],
    [0.15, 0.12, 0.13, 0.11],
    [0.1 , 0.1 , 0.1 , 0.12]
])

c_pm = -p
b_pm = np.concatenate((m, d))
A_pm = np.vstack((H, np.eye(len(p))))

cbar_pm, Abar_pm, b_pm = to_standard_form(c_pm, A_pm, b_pm)

# finally run the test!!
D_pm = build_tableau(cbar_pm, Abar_pm, b_pm)
assert D_pm.shape == (8, 12)
assert D_pm[0,0] == 0
assert np.allclose(D_pm.sum(axis=0)[6:], -np.ones(6))
assert np.allclose(D_pm[1:, 0] , b_pm)
```

**Exercise 3.2**

Implement the `should_terminate` function below

<!-- grade_id: ac6b53, points: 0, grade: false, solution: true -->
```python
def should_terminate(D: np.ndarray) -> bool:
    """
    Given a tableau D, check the termination condition which is
    that all of D[0, 1:] >= 0
    """
    return all(D[0, 1:]>= 0)
```

<!-- grade_id: ex3_2, points: 10, grade: true, solution: false -->
```python
assert should_terminate(np.array([[0, 1], [1, 0]]))
assert should_terminate(np.array([[0, 0], [1, 0]]))
assert not should_terminate(np.array([[0, -1], [1, 0]]))
```

**Exercise 3.3**

Implement the `pivot_column` function below

<!-- grade_id: fcee37, points: 0, grade: false, solution: true -->
```python
def pivot_column(D) -> int:
    """
    Given a tableau D, find the next pivot column

    If you cannot find another pivot column, raise a ValueError
    """
    A=0
    m,n=D.shape
    for i in range(1, n):
        if D[0,i]<0:
            A=i
            break
    if A==0:
        raise ValueError
    return A
```

<!-- grade_id: ex3_3, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_raises
assert_raises(ValueError, pivot_column, np.array([[-1, 0, 1], [0,0,0]]))
assert pivot_column(np.array([[-1, -1, 1], [0,0,0]])) == 1
assert pivot_column(np.array([[1, -1, 1], [0,0,0]])) == 1
assert pivot_column(np.array([[1, -1, -100], [0,0,0]])) == 1
assert pivot_column(np.array([[1, 0, -100], [0,0,0]])) == 2
```

#### 3.C: Find Pivot row

With the pivot column (j) in hand, we now need to find pivot row

To do this we follow these steps:

- Compute $d_i = - \frac{D_{i, 0}}{D_{i, j}}$ for all $i > 0$ (ignore the first row)
- Select the row $i$ whose value $d_i$ is the smallest


> Note: if two rows have same value for $d$, select the row with smaller index

> Note: If $D_{i, j} = 0$ for any $i$, we do not consider row $i$

**Exercise 3.4**

Implement the `pivot_row` function below

<!-- grade_id: 587c6e, points: 0, grade: false, solution: true -->
```python
def pivot_row(D, col: int) -> int:
    """
    Given a tableau D and pivot column `col`, find the next pivot column
    """
    d=[]
    m,n=D.shape
    for i in range(1,m):
        x= - D[i,0]/D[i,col]
        if x!=0:
            d.append(x)
    return d.index(min(d))+1
```

<!-- grade_id: ee1ad9, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
assert pivot_row(D_pm, 1) == 4
assert pivot_row(D_pm, 2) == 5
assert pivot_row(D_pm, 3) == 6
assert pivot_row(D_pm, 4) == 7
```

Output:
```
/var/folders/18/d01zbk1j4gn7w41cpr5tmfcr0000gn/T/ipykernel_963/103168367.py:8: RuntimeWarning: divide by zero encountered in double_scalars
  x= - D[i,0]/D[i,col]
```

#### 3.D: Do pivot operation

Now that we know the column (`j`) and row (`i`) for pivoting, we do the pivot operation

The goal of the pivot operation is to use the three matrix row operations to turn column `j` into the negative `i`th unit vector

That is we want column `j` to be all `0`, except for a `-1` at row `i`# [Modify the tests below for your own problem]

There is a strategy for doing this:

1. First multiply row `i` of `D` by `-1/D[i,j]`
2. For all rows `i2` where `i2 != i` and `D[i2, j] != 0`, multiply row `i` by `D[i2, j]` and add to row `i2` (this should set `D[i2, j]` to 0)

**Exercise 3.5**

Implement the `do_pivot` function below

<!-- grade_id: d1625e, points: 0, grade: false, solution: true -->
```python
def do_pivot(D, col, row) -> np.ndarray:
    """
    Given a tableau D and a pivot `col` and `row`, do the pivot operation
    """
    # apply step 1. of strategy to make D[row, col] = 1
    D[row, :]= D[row,:]*(-1/D[row, col]) # TODO: Fix this line

    for i2 in range(D.shape[0]):
        if i2 != row and D[i2, col]!= 0: 
            # your code here Apply step 2. of strategy above
            D[i2, :]+=D[row, :]*D[i2, col]  # TODO: Fix this line
    return D
```

<!-- grade_id: ex3_5, points: 10, grade: true, solution: false -->
```python
def make_negative_e(n, i):
    out = np.zeros(n)
    out[i] = -1
    return out

N = D_pm.shape[0]
for i in range(1,5):
    D_i = do_pivot(D_pm, i, i+3)
    assert np.allclose(D_i[:, i], make_negative_e(N, i+3))
```